---
title: 'Spreadsheet Foundations'
description: 'Chapter description goes here.'
---

## Understanding and Testing Data Types

```yaml
type: VideoExercise
key: e58fba4bd7
xp: 50
```

`@projector_key`
e51ae04f13cc2c4004e7fd2c960417b3

---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
